```python
# Initialize Otter
import otter
grader = otter.Notebook("lec_act_3_plotting.ipynb")
```

# Lecture goals

1. Understand the difference between a figure/window and a plotting area/subplot/axes
2. Be able to create a figure with subplots in it
3. Plot something in the subplot and label it
4. Pass parameters to functions by both position and name

Note: To see what the correct result should look like, see Lab lecture https://docs.google.com/presentation/d/1IiGGUNet-4Nj07x2cTXU6IOYXy9TSdAF5OUWCCKIYEM/edit?usp=sharing

# Install a pip package in the current Jupyter kernel
# Note: this is a safety check in case your JN install was not in the same place as your libraries
import sys
!{sys.executable} -m pip install matplotlib

```python
# Access all numpy functions as np.
import numpy as np
import matplotlib.pyplot as plt
%matplotlib inline
import json as json
```

## Read data

TODO: Read in the Wrist_torque_data.csv file. This is the data we'll be plotting; note that this data is the result of calling **get_channel_data** for the wrist torque data - it is **n_picks** by **n_timesteps * 3**.

```python
# Load in the Wrist_torque_data.csv file
#  Note: Use the filename "Data/Wrist_torque_data.csv" - do not put in a full path name/a different file name, or the
#   autograder on Gradescope will fail and you'll get a zero.
file_name = "Data/Wrist_torque_data.csv"
pick_data = np.loadtxt(file_name, dtype ="float", delimiter = ',')

wrist_torque_data = pick_data
```

```python
grader.check("Read data")
```

Output:
```
Read data results: All test cases passed!
```

## Get data

TODO: Get the Wrist torque **z** channel for the first row of pick data (these will be the y values you plot with). Create the **t** (time) values (these will be the x values you plot with). Remember that the z values start at index 2 and are every 3rd value

For the **t** values, assume the data is sampled at 30 Hz, i.e., the time sampling is 1/30th of a second

- Step 1: How big does the t array have to be? (hint: How many z values are there for the wrist torque data? What was n_time_steps?)
- Step 2: How do you make an array of that size with that spacing? (hint: **np.arange**)


```python
# TODO: Create a numpy array that starts at 0, ends at number of time steps * time_step, and has step size time_step
#   See np.arange
# Step 1: Calculate the number of time steps in the data (assume a data spacing of 3)


# Step 2: Create a numpy array with that many elements with a spacing of 1.0/30.0 (30 Hz) (this is ts)

# Step 3: Get out just the z values from the wrist torque data (this is z_values) for the first row


# Print the results



# TODO Set up the time (x axis) data 
time_step = 1.0 / 30.0  # This is 1/30th of a second
n_time_steps = wrist_torque_data.shape[1] // 3
ts = np.arange(0, n_time_steps * time_step, time_step)
z_values = wrist_torque_data[0][2::3]

# Print the results
print("Time Steps Array:", ts)
print("Z Values for the First Row:", z_values)

# Use this to check that your data is the right size - both should be of size (40,)
#  Note: it says (40,) instead of (40,1). That's just numpy's way of saying that it's a 1 dimensional array,
#   or a tuple with 1 value
print(f"ts is size {ts.shape}")
print(f"z_values is size {z_values.shape}")

```

Output:
```
Time Steps Array: [0.         0.03333333 0.06666667 0.1        0.13333333 0.16666667
 0.2        0.23333333 0.26666667 0.3        0.33333333 0.36666667
 0.4        0.43333333 0.46666667 0.5        0.53333333 0.56666667
 0.6        0.63333333 0.66666667 0.7        0.73333333 0.76666667
 0.8        0.83333333 0.86666667 0.9        0.93333333 0.96666667
 1.         1.03333333 1.06666667 1.1        1.13333333 1.16666667
 1.2        1.23333333 1.26666667 1.3       ]
Z Values for the First Row: [ 0.1285  0.1286  0.1266  0.1273  0.126   0.1257  0.126   0.1238  0.1242
  0.1292  0.1274  0.1342  0.1326  0.1316  0.1316  0.131   0.1305  0.1299
  0.1295  0.1295  0.1228  0.1227  0.1228  0.1216  0.1151 -0.0729 -0.1074
 -0.0916 -0.0772 -0.0735 -0.0832 -0.0973 -0.106  -0.0883 -0.0806 -0.1006
 -0.1025 -0.0885 -0.0831 -0.0773]
ts is size (40,)
z_values is size (40,)
```

```python
grader.check("Get data")
```

Output:
```
Get data results: All test cases passed!
```

## Plot data

TODO: Plot the **ts** versus **z_values** data. Make sure to add x and y labels, along with a title.

x-axis label: time - make sure to use units (seconds) in the label

y-axis label: what you're plotting - make sure you use the units (N.m)

Title: Should be what data channel this is

```python
# Create the plotting window. Use subplots.
nrows = 1
ncols = 1


# Set the number of rows and columns for subplots
nrows = 1
ncols = 1

# Create the plotting window with subplots
fig, axs = plt.subplots(nrows=nrows, ncols=ncols, subplot_kw={'projection': '2d'})

# Extract data for the first row
x = df.iloc[0, 3]  # Torque x
y = df.iloc[0, 4]  # Torque y
z = df.iloc[0, 5]  # Z-axis data

# Create a 3D scatter plot
axs.scatter(x, y, z)

# Customize the plot if needed
axs.set_xlabel('Torque X')
axs.set_ylabel('Torque Y')
axs.set_zlabel('Z-axis Data')

# Show the plot
plt.show()


plt.show()
```

Output:
```
File c:\Users\gatis\anaconda3\Lib\site-packages\matplotlib\figure.py:1721, in FigureBase._process_projection_requirements(self, axes_class, polar, projection, *args, **kwargs)
   1718     projection = 'polar'
   1720 if isinstance(projection, str) or projection is None:
-> 1721     projection_class = projections.get_projection_class(projection)
   1722 elif hasattr(projection, '_as_mpl_axes'):
   1723     projection_class, extra_kwargs = projection._as_mpl_axes()

File c:\Users\gatis\anaconda3\Lib\site-packages\matplotlib\projections\__init__.py:110, in get_projection_class(projection)
    108     return projection_registry.get_projection_class(projection)
    109 except KeyError as err:
--> 110     raise ValueError("Unknown projection %r" % projection) from err

ValueError: Unknown projection '2d'
<Figure size 640x480 with 0 Axes>
```

```python
grader.check("Plot data")
```

Output:
```
Plot data results:
    Plot data - 1 result:
        ✅ Test case passed

    Plot data - 2 result:
        ❌ Test case failed
        Trying:
            assert("s" in axs.get_xlabel() or "S" in axs.get_xlabel())
        Expecting nothing
        **********************************************************************
        Line 1, in Plot data 1
        Failed example:
            assert("s" in axs.get_xlabel() or "S" in axs.get_xlabel())
        Exception raised:
            Traceback (most recent call last):
              File "c:\Users\gatis\anaconda3\Lib\doctest.py", line 1351, in __run
                exec(compile(example.source, filename, "single",
              File "<doctest Plot data 1[0]>", line 1, in <module>
                assert("s" in axs.get_xlabel() or "S" in axs.get_xlabel())
                                                  ^^^^^^^^^^^^^^^^^^^^^^^
            AssertionError

    Plot data - 3 result:
        ❌ Test case failed
        Trying:
            assert("N.m
```

## Hours and collaborators
Required for every assignment - fill out before you hand-in.

Listing names and websites helps you to document who you worked with and what internet help you received in the case of any plagiarism issues. You should list names of anyone (in class or not) who has substantially helped you with an assignment - or anyone you have *helped*. You do not need to list TAs.

Listing hours helps us track if the assignments are too long.

```python

# List of names (creates a set)
worked_with_names = {"not filled out"}
# List of URLS TCW3 (creates a set)
websites = {"not filled out"}
# Approximate number of hours, including lab/in-class time
hours = -1.5
```

```python
grader.check("hours_collaborators")
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

Submit through gradescope, Lecture activity 3 plotting. Submit just the .ipynb file - the data files will be included for you

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False, run_tests=True)
```

Tests for Get data (points: 2)
```
>>> assert(ts.size == 40)
```
```
>>> assert(z_values.size == 40)
```
```
>>> assert(np.isclose(ts[0], 0.0))
```
```
>>> assert(np.isclose(ts[-1], 1.3))
```
```
>>> assert(np.isclose(z_values[0], 0.12849379, atol=0.001))
```
```
>>> assert(np.isclose(z_values[-1], -0.077284, atol=0.001))
```


Tests for Plot data (points: 2)
```
>>> assert(ts.shape == z_values.shape)
```
```
>>> assert("s" in axs.get_xlabel() or "S" in axs.get_xlabel())
```
```
>>> assert("N.m" in axs.get_ylabel())
```
```
>>> assert("rist" in axs.get_title() and "orque" in axs.get_title())
```


Tests for Read data (points: 1)
```
>>> assert(wrist_torque_data.shape == (660, 40 * 3))
```


Tests for hours_collaborators (points: 1)
```
>>> assert(not "not filled out" in worked_with_names)
```
```
>>> assert(not "not filled out" in websites)
```
```
>>> assert(hours > 0)
```
